<font size="4" style="color:red;"> **IMPORTANT: ** When submitting this homework notebook, please modify only the cells that start with:</font>

```python
# modify this cell
```

### Import Stuff

Notice that we do not import *numpy* or *scipy* neither of these packages are need for this homework. For our solutions, the only command we needed to import was `itertools.product()`

```python
import itertools
from itertools import product
```

# Sets
Read the notebook on sets before attempting these exercises

# Problem 1

De Morgan's first law states the following for any two sets $A$ and $B$
$$(A\cup B)^c = A^c\cap B^c$$

In the following two exercises we calculate $(A\cup B)^c$ in two different ways. Both functions must take $A$, $B$ and the universal set $U$ as their inputs.

## Exercise 1.1


Write the function **complement_of_union** that first determines $A\cup B$ and then evaluates the complement of this set. Output the tuple: $\begin{pmatrix}A\cup B,\, (A\cup B)^c\end{pmatrix}$.



<font  style="color:blue"> **Code**</font>
```python
A = {1, 2, 3}
B = {3, -6, 2, 0}
U = {-10, -9, -8, -7, -6, 0, 1, 2, 3, 4}
complement_of_union(A, B, U)
```

<font  style="color:magenta"> **Output**</font>
```
({-6, 0, 1, 2, 3}, {-10, -9, -8, -7, 4})
```


```python
# modify this cell

def complement_of_union(A, B, U):
    # inputs: A, B and U are of type 'set'
    # output: a tuple of the type (set, set)
    unionAB= A.union(B)
    complimentAB= U.difference(unionAB)
    return (unionAB, complimentAB)

```

<!-- grade_id: ex1, points: 5, grade: true, solution: false -->
```python
# Check Function

A = {-6, 3, 4, 5}
B = {-6, 5, 13}
U = A|B|{12,-2, -4}
x=complement_of_union(A,B,U)
print(x)
#assert( complement_of_union(A, B, U) == ({0, 1, 2, 3, 4, 5, 7, 8, 9, 10, -6, -4, -3}, {-4, -3, 7, 10})  )

#
# AUTOGRADER TEST - DO NOT REMOVE
#

```

Output:
```
({3, 4, 5, -6, 13}, {-4, 12, -2})
```

## Exercsise 1.2

Write the function **intersection_of_complements** that first determines $A^c$ and $B^c$ and then evaluates the intersection of their complements. Output the tuple: $\begin{pmatrix}A^c, \,  A^c\cap B^c\end{pmatrix}$

<font  style="color:blue"> **Code**</font>
```python
A = {1, 2, 3}
B = {3, -6, 2, 0}
U = {-10, -9, -8, -7, -6, 0, 1, 2, 3, 4}
intersection_of_complements(A, B, U)
```

<font  style="color:magenta"> **Output**</font>
```
({-10, -9, -8, -7, -6, 0, 4}, {-10, -9, -8, -7, 4})
```


```python
# modify this cell

def intersection_of_complements(A, B, U):
    # inputs: A, B and U are of type 'set'
    # output: a tuple of the form (set, set)
    Ac= U-A
    AcBc= Ac.intersection(U-B)
    return(Ac,AcBc)

```

<!-- grade_id: ex2, points: 5, grade: true, solution: false -->
```python
# Check Function

A = {-6, 3, 4, 5}
B = {-6, 5, 13}
U = A|B|{-2, 12, -4}
#assert(  intersection_of_complements(A, B, U) == ({-6, -4, -3, 0, 7, 8, 9, 10}, {-4, -3, 7, 10})  )
x= intersection_of_complements(A, B, U)
print(x)
#
# AUTOGRADER TEST - DO NOT REMOVE
#

```

Output:
```
({13, -4, 12, -2}, {-4, 12, -2})
```

# Problem 2

This problem illustrates a property of cartesian products of unions of two or more sets. For four sets $A$, $B$, $S$ and $T$, the following holds:

$$(A\cup B)\times(S\cup T) = (A\times S)\cup(A\times T)\cup(B\times S)\cup(B\times T)$$

Write the following functions to determine $(A\cup B)\times(S\cup T)$ in two different ways.


## Exercies 2.1

Write function **product_of_unions** that first determines $(A\cup B)$ and $(S\cup T)$ and then evaluates the cartesian products of these unions. Output the tuple $\begin{pmatrix}(A\cup B),\,  (A\cup B)\times(S\cup T)\end{pmatrix}$.

<font  style="color:blue"> **Code**</font>
```python
A = {1, 2}
B = {1, 3}
S = {-1, 0}
T = {0, 10}
product_of_unions(A, B, S, T)
```


<font  style="color:magenta"> **Output**</font>
```
({1, 2, 3},
 {(1, -1),
  (1, 0),
  (1, 10),
  (2, -1),
  (2, 0),
  (2, 10),
  (3, -1),
  (3, 0),
  (3, 10)})
```


```python
# modify this cell

def product_of_unions(A, B, S, T):
    # inputs: A, B, S and T are sets
    # output: a tuple of the type (set, set)
    unionAB= A.union(B)
    unionST= S.union(T)
    x=set()
    for pair in product(unionAB, unionST):
        x.add(pair)
    return(unionAB, x)


```

<!-- grade_id: ex3, points: 5, grade: true, solution: false -->
```python
# Check Function

A = {5}
B = {5}
S = {-1, 0}
T = {0}
#assert( product_of_unions(A, B, S, T) == ({5, 6}, {(5, -1), (5, 0), (5, 1), (5, 2), (6, -1), (6, 0), (6, 1), (6, 2)})   )
y= product_of_unions(A, B, S, T)
print(y)
#
# AUTOGRADER TEST - DO NOT REMOVE
#

```

Output:
```
({5}, {(5, -1), (5, 0)})
```

## Exercise 2.2

Write a function **union_of_products** that first determines $(A\times S)$ and the other three cartesian products that appear on the right hand side of the identity above, then evaluates the union of these cartesian products. Output the tuple $\begin{pmatrix}(A\times S),\,  (A\times S)\cup(A\times T)\cup(B\times S)\cup(B\times T)\end{pmatrix}$.

<font  style="color:blue"> **Code**</font>
```python
A = {1, 2}
B = {1, 3}
S = {-1, 0}
T = {0, 10}
union_of_products(A, B, S, T)
```


<font  style="color:magenta"> **Output**</font>
```
({(1, -1), (1, 0), (2, -1), (2, 0)},
 {(1, -1),
  (1, 0),
  (1, 10),
  (2, -1),
  (2, 0),
  (2, 10),
  (3, -1),
  (3, 0),
  (3, 10)})
```

```python
# modify this cell

def union_of_products(A, B, S, T):
    # inputs: A, B, S and T are sets
    # output: a tuple of the type (set, set)
    productAS= set()
    for pair in product(A,S):
        productAS.add(pair)
    productAT= set()
    for pair in product(A,T):
        productAT.add(pair)
    productBS= set()
    for pair in product(B,S):
        productBS.add(pair)
    productBT= set()
    for pair in product(B,T):
        productBT.add(pair)
    unionfinal= productAS.union(productAT.union(productBS.union(productBT)))
    return(productAS, unionfinal)

```

<!-- grade_id: ex4, points: 5, grade: true, solution: false -->
```python
# Check Function

A = {5}
B = {5}
S = {-1, 0}
T = {0}
#assert( union_of_products(A, B, S, T) == \
#        ({(5, -1), (5, 0), (5, 1)}, \
 #        {(5, -1), (5, 0), (5, 1), (5, 2), (6, -1), (6, 0), (6, 1), (6, 2)})  \
  #    )
z= union_of_products(A, B, S, T)
print(z)
#
# AUTOGRADER TEST - DO NOT REMOVE
#

```

Output:
```
({(5, -1), (5, 0)}, {(5, -1), (5, 0)})
```

```python

```